Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Qinwen Zheng"
COLLABORATORS = ""
```

---

# Homework 5

**COMPSCI 101, Duke Kunshan University**

# Problem 1 (10 point) 
We have 13 lamps, labeled 1 through 13, arranged in a sequence. Starting with lamp 1, we count each lamp that is still on. Every time we reach the 5th lamp in our counting, we turn that lamp off. Once we reach lamp 13, if there are still lamps on, we cycle back to lamp 1 and continue the process. We do not count any lamps that have already been turned off. This continues until all lamps are turned off.

Task: Write a function that returns the order in which the lamps are turned off. For example, if we apply this procedure to 13 lamps with a step of 5, the resulting turn-off order is: [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6].

<!-- grade_id: cell-06952d5f38104138, points: 0, grade: false, solution: true -->
```python
def turn_off_lamps_order():
    num_lamps = 13
    start = 0
    step = 5
    res = []
    lights = []
    for i in range(num_lamps):
        lights.append(1)
    
    counter = 1 # count the step
    ind = start # the index 
    c = 0 #how many lighs are turned off
    while True:
        if c == num_lamps:
            break
        else:
            if lights[ind] == 0:
                if ind == num_lamps - 1:
                    ind = 0
                else:
                    ind +=1
            else:
                if counter == step:
                    lights[ind] = 0
                    res.append(ind + 1)
                    counter = 1
                    if ind == num_lamps - 1:
                        ind = 0
                    else:
                        ind += 1
                    
                    c += 1
                else:
                    counter += 1
                    if ind == num_lamps - 1:
                        ind = 0
                    else:
                        ind +=1
    return res

```

<!-- grade_id: cell-bf13653d315bd398, points: 10, grade: true, solution: false -->
```python
assert turn_off_lamps_order() == [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6]
```

# Problem 2 (10 point) 
According to problem 1, what if we have different number of lamps, and start at a lamp other than the first lamp. Even more, when we turn off the lamp at every kth step. Write a function that require 3 numbers, representing the number of lamps, start from which lamp, every how many steps to turn off a lamp, and return the turn-off order.

<!-- grade_id: cell-b47caf9d62efd4e1, points: 0, grade: false, solution: true -->
```python
def turn_off_order_upgrade(num_lamps=13, start=0, step=5):
    """
    Returns the order in which lamps are turned off,
    turning off every 'step'-th lamp among those still active.
    """
    res = []
    lights = []
    for i in range(num_lamps):
        lights.append(1)
    
    counter = 1 # count the step
    ind = start # the index 
    c = 0 #how many lighs are turned off
    while True:
        if c == num_lamps:
            break
        else:
            if lights[ind] == 0:
                if ind == num_lamps - 1:
                    ind = 0
                else:
                    ind +=1
            else:
                if counter == step:
                    lights[ind] = 0
                    res.append(ind + 1)
                    counter = 1
                    if ind == num_lamps - 1:
                        ind = 0
                    else:
                        ind += 1
                    
                    c += 1
                else:
                    counter += 1
                    if ind == num_lamps - 1:
                        ind = 0
                    else:
                        ind +=1
    return res
```

<!-- grade_id: cell-aac34bc3384ca811, points: 10, grade: true, solution: false -->
```python
# hidden tests in this cell
```

# Problem 3 (10 points) 

**function**: `move_zero`

**input**: a list

**output**: a list

Write a function called `move_zero` that that takes a list of integers, and then returns a new list which move all zero digits to the end of the input list.

Example:
[3, 4, 0, 0, 0, 6, 2, 0, 6, 7, 6, 0, 0, 0, 9, 10, 7, 4, 4, 5, 3, 0, 0, 2, 9, 7, 1]

Move all zero digits to end of the said list of numbers:

[3, 4, 6, 2, 6, 7, 6, 9, 10, 7, 4, 4, 5, 3, 2, 9, 7, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0]

<!-- grade_id: problem3_example, points: 0, grade: false, solution: true -->
```python
def move_zero(input_list):
    res = []
    counter = 0
    for i in range (len(input_list)):
        if input_list[i] != 0:
            res.append(input_list[i])
            counter += 1
    for i in range(len(input_list)- counter):
        res.append(0)
    return res
```

<!-- grade_id: cell-413652036a3b4661, points: 10, grade: true, solution: false -->
```python
assert move_zero([3, 4, 0, 0, 0, 6, 2, 0, 6, 7, 6, 0, 0, 0, 9, 10, 7, 4, 4, 5, 3, 0, 0, 2, 9, 7, 1]) == [3, 4, 6, 2, 6, 7, 6, 9, 10, 7, 4, 4, 5, 3, 2, 9, 7, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0]
# hidden tests in this cell
```

# Problem 4 (10 point)

One apple is 5 dollars, one pear is 3 dollars, and 3 bananas are 1 dollar. How to buy 100 fruits using 100 dollars?

Input: no input

Output: all possible combonation of `[#apples, #pears, #bananas]`, the output should be a list of list, e.g. `[[#apples, #pears, #bananas], [#apples, #pears, #bananas], [#apples, #pears, #bananas]]`


<!-- grade_id: cell-0b95d1b83eb37a8c, points: 0, grade: false, solution: true -->
```python
def fun():
    results = []
    for a in range(101):  
        for p in range(101 - a):
            b = 100 - a - p 
            if 5 * a + 3 * p + b / 3 == 100:
                results.append([a, p, b])
    return results
```

Output:
```
[[0, 25, 75], [4, 18, 78], [8, 11, 81], [12, 4, 84]]
```

<!-- grade_id: cell-9f5764dd7239884e, points: 10, grade: true, solution: false -->
```python
# hidden tests in this cell

```

# Problem 5 (10 point)

Write a function called `get_indices` that takes a string `text` as parameter. It returns a list of  the indices of all non-vowels in the order they appear in the `text`, excluding the space character. You may assume the `text` contains only the alphabet letters and space. For example, `find_indices('I love you')` will return the list [2, 4, 7].

<!-- grade_id: cell-02c79d88e0f65e6e, points: 0, grade: false, solution: true -->
```python
def get_indices(text):
    vowels = 'aeiouAEIOU'
    indices = []
    for i, ch in enumerate(text):
        if ch not in vowels and ch != ' ':
            indices.append(i + 1)  
    return indices
```

<!-- grade_id: cell-6da4d6fe553f2c3a, points: 10, grade: true, solution: false -->
```python
# hidden tests in this cell

```

# Problem 6 (10 point)

You are given an array a of integers, where each element satisfies  1 ≤ a[i] ≤ n (with  n  being the size of the array). Some elements in the array may appear more than once, while others may be missing. Your task is to find all the numbers in the range [1, n] that are missing from the array.
- If  n < 1 , return an empty list [].
- If all numbers in [1, n] are present in the array, return an empty list [].

Examples:

Input: [4, 3, 2, 7, 8, 2, 3, 1]
Output: [5, 6]

Input: [1, 1]
Output: [2]

Input: [1, 2, 3, 4]
Output: []

<!-- grade_id: cell-82e2f5928a22758d, points: 0, grade: false, solution: true -->
```python
def findDisappearedNumbers(nums):
    """
    Find all numbers in the range [1, n] that are missing from the array nums.

    :param nums: List of integers
    :return: List of missing integers
    """
    n = len(a)
    if n < 1:
        return []
    for num in a:
        index = abs(num) - 1
        if a[index] > 0:
            a[index] = -a[index]

    missing = [i + 1 for i in range(n) if a[i] > 0]
    return missing
```

<!-- grade_id: cell-447d90122b0aaf79, points: 10, grade: true, solution: false -->
```python
# hidden tests in this cell

```

# Problem 7 (10 point)


Given a list of integers, return indices of the two numbers such that they add up to a specific target.

Indices should be in an increasing order

You may assume that each input would have exactly one solution, and you may not use the same element twice.

<!-- grade_id: cell-f9fb4ca484030688, points: 0, grade: false, solution: true -->
```python
def TwoSum(nums, target):
    lookup = {}
    for i, num in enumerate(nums):
        complement = target - num
        if complement in lookup:
            return sorted([lookup[complement], i])
        lookup[num] = i
```

<!-- grade_id: cell-82183aa329825131, points: 10, grade: true, solution: false -->
```python
assert (TwoSum([2, 7, 11, 15],9)==[0,1])
# hidden tests in this cell

```