Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Qinwen Zheng"
COLLABORATORS = ""
```

---

# Homework 04
### COMPSCI 101, Duke Kunshan University

#### Author: Tongshu Zheng

# Problem 1 (10 point) 
Given a 32-bit signed integer, reverse digits of an integer.
If the return number is outside this range \\( [-2^{31},2^{31}-1] \\), then return 0

example:  input 12345  output 54321

<!-- grade_id: reverse_integer, points: 0, grade: false, solution: true -->
```python
def reverse_integer(x):
    ''':type x: int
       :rtype: int
    '''

    is_minor = False
    if x < 0:
        is_minor = True
        x = x * (-1)

    xstr = str(x)
    res = ""
    for i in range(len(xstr)-1 , -1, -1):
        res = res + xstr[i]

    ress = 0

    if is_minor == True:
        res = '-' + res
        ress = int(res)
    else:
        ress = int(res)

    if ress < -(2**31) or ress > 2**31 - 1:
        return 0
    else:
        return ress
```

<!-- grade_id: test_myreverse, points: 10, grade: true, solution: false -->
```python
def test_myreverse():
    # Test 1: Positive number
    assert reverse_integer(123) == 321, "Test Case 1 Failed"
    
    # Test 2: Negative number
    assert reverse_integer(-123) == -321, "Test Case 2 Failed"
    
    # Test 3: Single-digit positive number
    assert reverse_integer(7) == 7, "Test Case 3 Failed"
    
    # Test 4: Single-digit negative number
    assert reverse_integer(-5) == -5, "Test Case 4 Failed"
    
    # Test 5: Number with trailing zeros
    assert reverse_integer(120) == 21, "Test Case 5 Failed"
    
    # Test 6: Zero input
    assert reverse_integer(0) == 0, "Test Case 6 Failed"
    
    # Test 7: Large positive number within 32-bit range
    assert reverse_integer(1463847412) == 2147483641, "Test Case 7 Failed"
    
    # Test 8: Large negative number within 32-bit range
    assert reverse_integer(-2147483412) == -2143847412, "Test Case 8 Failed"
    
    # Test 9: Positive number that overflows
    assert reverse_integer(1534236469) == 0, "Test Case 9 Failed"
    
    # Test 10: Negative number that overflows
    assert reverse_integer(-1534236469) == 0, "Test Case 10 Failed"
    
    # Test 11: Edge case for maximum 32-bit positive integer
    assert reverse_integer(2147483647) == 0, "Test Case 11 Failed"  # Reversed value exceeds 32-bit range
    
    # Test 12: Edge case for minimum 32-bit negative integer
    assert reverse_integer(-2147483648) == 0, "Test Case 12 Failed"  # Reversed value exceeds 32-bit range
    

    print("All test cases passed!")
    
test_myreverse()
```

Output:
```
Cell In[3], line 41
     36     assert reverse_integer(-2147483648) == 0, "Test Case 12 Failed"  # Reversed value exceeds 32-bit range
     39     print("All test cases passed!")
---> 41 test_myreverse()

Cell In[3], line 3, in test_myreverse()
      1 def test_myreverse():
      2     # Test 1: Positive number
----> 3     assert reverse_integer(123) == 321, "Test Case 1 Failed"
      5     # Test 2: Negative number
      6     assert reverse_integer(-123) == -321, "Test Case 2 Failed"

AssertionError: Test Case 1 Failed
```

# Problem 2 (10 point) 
Write a function to find the longest common prefix string amongst a list of strings.

If there is no common prefix, return an empty string ''.

example: input ['what','where','who']     output 'wh'

<!-- grade_id: longestCommonPrefix, points: 0, grade: false, solution: true -->
```python
def longest_common_prefix(strs):
    """
    Find the longest common !!case-sensitive!! prefix string among a list of strings.
    
    :param strs: List of strings
    :return: The longest common prefix string, or an empty string '' if none exists
    """
    if not strs:
        return ""
    
    prefix = strs[0]
    
    for s in strs[1:]:
        while not s.startswith(prefix):
            prefix = prefix[:-1]
            if not prefix:
                return ""
    
    return prefix
```

<!-- grade_id: longestCommonPrefix_test, points: 10, grade: true, solution: false -->
```python
def test_longest_common_prefix():
    # 1. General cases
    assert longest_common_prefix(["what", "where", "who"]) == "wh", "Test Case 1 Failed"
    assert longest_common_prefix(["flower", "flow", "flight"]) == "fl", "Test Case 2 Failed"
    
    # 2. No common prefix
    assert longest_common_prefix(["dog", "racecar", "car"]) == "", "Test Case 3 Failed"
    assert longest_common_prefix(["apple", "banana", "cherry"]) == "", "Test Case 4 Failed"
    
    # 3. Single string
    assert longest_common_prefix(["single"]) == "single", "Test Case 5 Failed"
    
    
    
    print("All test cases passed!")

test_longest_common_prefix()
```

# Problem 3 (5 point) 

Write a function isPalindrome(s) that returns True if the string s is a palindrome, and False otherwise. A palindrome is a string that reads the same forwards and backwards.

For example:

'A man a plan a canal Panama' is palindrome

'No lemon no melon' is palindrome

'Was it a car or a cat I saw' is palindrome


<!-- grade_id: cell-e00f009e6ec3b5c1, points: 0, grade: false, solution: true -->
```python
def isPalindrome(word):
    if word == "":
        return True

    word = word.replace(" ", "")
    word = word.lower()
    
    for i in range(0, len(word)):
        if word[i] != word[len(word) - i - 1]:
            return False
    
    return True
```

<!-- grade_id: cell-26f4b8de7c54e195, points: 5, grade: true, solution: false -->
```python
assert isPalindrome("A man a plan a canal Panama") == True
assert isPalindrome("No lemon no melon") == True
assert isPalindrome("Was it a car or a cat I saw") == True
assert isPalindrome("a") == True
assert isPalindrome(" ") == True
assert isPalindrome("python") == False
```

# Problem 4 (10 point) 

Develop a function `longestPalindrome(seq)` that takes a sequence `seq` and returns the longest palindrome in the sequence. For example, the input ‘mdadc’ gives ‘dad’, and the input ‘rtyc456u654cymnredder’ gives ‘yc456u654cy’. 

**Note**: you should use isPalindrome(word) from Problem 3 to solve this problem.

<!-- grade_id: cell-844e58d672fd97e5, points: 0, grade: false, solution: true -->
```python
def longestPalindrome(seq):
    
    def isPalindrome(word):
        if word == "":
            return True

        word = word.replace(" ", "")
        word = word.lower()
        
        for i in range(0, len(word)):
            if word[i] != word[len(word) - i - 1]:
                return False
        
        return True

    res = ""
    for sta in range( len(seq) ):
        for end in range( sta, len(seq) ):
            if isPalindrome( seq[sta:end] ) == True:
                if end - sta > len(res):
                    res = seq[sta:end]

    return res

```

<!-- grade_id: cell-ff9f6efe082c0385, points: 10, grade: true, solution: false -->
```python
assert longestPalindrome('rtyc456u654cymnredder') == 'yc456u654cy'
assert longestPalindrome("abcbaHELLOmadamTEST") == 'abcba'
assert longestPalindrome("bc") == 'b' or longestPalindrome("bc") == 'c'
```