Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Qinwen Zheng"
COLLABORATORS = ""
```

---

# Homework 7
**COMPSCI 101, Duke Kunshan University**

## Problem 1

**function**: `list_sum_recur`

**input**: `num_List` (whose type is list)

**output**: list_sum (whose type is integer)

Write a function `list_sum_recur` using recursion to calculate the sum of a list of numbers.
For example: input parameter list [2, 5, 7, 8, -10]
Return: 12

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def list_sum_recur(num_List):
    if len(num_List) == 1:
        return num_List[0]
    else:
        return num_List[0] + list_sum_recur(num_List[1:])
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 10, grade: true, solution: false -->
```python
assert list_sum_recur([2,4,5,6,7]) == 24
# Hidden tests here
```

## Problem 2

**function**: `sum_recur`

**input**: `n` (whose type is positive integer)

**output**:  a variable (whose type is float)

Write a function `sum_recur` using recursion that takes a positive integer n to calculate the following expression: 
1-1/2 +1/3 -1/4 +1/5 …… (hint you need check n is even or odd)

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def sum_recur(n):
    if n == 1:
        return 1
    else:
        if n%2 == 0:
            return sum_recur(n-1) - (1/n)
        else:
            return sum_recur(n-1) + (1/n)
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 10, grade: true, solution: false -->
```python
assert abs(sum_recur(5) - 0.7833333333333333) < 1e-10
# Hidden tests here
```

## Problem 3

**function**: `sum_dig_recur`

**input**: `n` (whose type is positive integer)

**output**:  a variable (whose type is integer)

Write a function `sum_dig_recur` using recursion that takes a positive integer `n` to calculate the sum of its digits.
Example: sum_dig_recur(345) -> 12

<!-- grade_id: cell-09e5f93a7fb96627, points: 0, grade: false, solution: true -->
```python
def sum_dig_recur(n):
    if len(str(n)) == 1 :
        return n
    else:
        strn = str(n)
        return int(strn[0]) + sum_dig_recur(int(strn[1:]))
```

<!-- grade_id: cell-f427fc8c08ea6cb1, points: 10, grade: true, solution: false -->
```python
assert sum_dig_recur(12345) == 15
# Hidden tests here
```

## Problem 4

Function: ```sum_series_recur```

input: an int

output: an int

Write a function “sum_series_recur” using recursion that takes a positive integer n to calculate the sum of the positive integers of n+(n-2)+(n-4) + (n - 6)... (until n-x =< 0).

Example: sum_series_recur(10) -> 30

<!-- grade_id: cell-b0f5d5ea7a6eb563, points: 0, grade: false, solution: true -->
```python
def sum_series_recur(n):
    if n <= 0 :
        return 0
    else:
        return sum_series_recur(n-2) + n
```

<!-- grade_id: cell-2e652d8d7a3a2f67, points: 10, grade: true, solution: false -->
```python
assert sum_series_recur(10) == 30
# Hidden tests here
```

## Problem 5
It is time for Jim to go home but the elevator is out of order. Jim has to climb a staircase with `n` steps. Each time he can either climb one or two steps.

Write a function called `num_ways` that takes an integer `n` as input and returns the number of distinct ways he reach the top of a staircase with `n` steps.



<!-- grade_id: cell-eb9f4e347802869a, points: 0, grade: false, solution: true -->
```python
def num_ways(n):
    if n == 1 or n == 0:
        return 1
    else:
        return num_ways(n-1) + num_ways(n-2)
```

<!-- grade_id: cell-f860fb8aeb7243c0, points: 10, grade: true, solution: false -->
```python
assert num_ways(2) == 2
assert num_ways(5) == 8

# Hidden tests here
```

## Problem 6


Write a recursive function `NumOfSubsets` that takes two arguments: an iterable data `S` (such as string, list, and tuple) and an integer `n`. It returns the number of subsets that can be formed from the elements in `S` and there are `n` elements in each subset. This must be solved recursively!


**Hint**:

The recursive solution is based on the combinatorial principle that:

The number of subsets of size n from set S equals:

The number of subsets that include the first element (so we need to choose n-1 more elements from the remaining elements)

Plus the number of subsets that don't include the first element (so we need to choose all n elements from the remaining elements)

Mathematically, this is represented by the combination formula: C(k, n) = C(k-1, n-1) + C(k-1, n)

<!-- grade_id: cell-9c4a9eb8acd923bc, points: 0, grade: false, solution: true -->
```python
def NumOfSubsets(S, n):
    if n == 0:
        return 1
    elif len(S) < n:
        return 0
    else:
        return NumOfSubsets(S[1:], n-1) + NumOfSubsets(S[1:], n)
```

<!-- grade_id: cell-c249dfb4ea02b3e7, points: 10, grade: true, solution: false -->
```python
assert NumOfSubsets([1, 2, 3], 2) == 3
# Hidden tests here
```

## Problem 7

Given a string containing just the characters `'('`, `')'`, `'{'`, `'}'`, `'['` and `']'`, use `recursion` to determine whether the input string is a complete parenthesis combination. An empty string should also be considered valid.

Return True if the string is a valid combination of parentheses, and False otherwise.

Some correct cases:
`{[()]}`, `[]{}()`, `[()]`, `{}` ...

<!-- grade_id: cell-069fef43cf9adec1, points: 0, grade: false, solution: true -->
```python
def isValidRecursive(s: str) -> bool:
    if s == "":
        return True
    else:
        for pair in ["()", "[]", "{}"]:
            if pair in s:
                return isValidRecursive(s.replace(pair, "", 1))
    return False
```

Output:
```
Running cells with '.venv (Python 3.14.0)' requires the ipykernel package.

Install 'ipykernel' into the Python environment. 

Command: '/Users/zhengqinwen/Documents/GitHub/CS101Codes/.venv/bin/python -m pip install ipykernel -U --force-reinstall'
```

<!-- grade_id: cell-6d107b6ac641e104, points: 10, grade: true, solution: false -->
```python
"""Check that isValid function returns the correct output"""
assert (isValidRecursive("[[[]]{}}") == False)
assert (isValidRecursive("[](){}{[]}()[[]]") == True)
# Hidden tests here
```