Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Qinwen Zheng"
COLLABORATORS = ""
```

---

# Homework 6
**COMPSCI 101, Duke Kunshan University**

## Author: Tongshu Zheng

## Problem 1:  Book Cipher 

### Part A: Generate Encoding Keys (5 points)

A book cipher is a cipher for which the key is derived from a book. For example, it might map each character in the plain text to the numeric index **as string** of the first occurrence of that character in the book (or on a page of the book). The assumption is that the sender and receiver of the coded message have previously agreed on the book, but an adversary who intercepts the coded message does not know what book or the page of the book was used to code it.

For example: If plain_text were "no is no" and book started with "Once upon a time, in a house in a land far away," the function `gen_code_keys(book, plain_text)` would return {'n': '1', 'o': '7', ' ': '4', 'i': '13', 's': '26'}. Notice, by the way, that o gets mapped to seven rather than zero because o and O are different characters.

<!-- grade_id: cell-f74115096cd90d09, points: 0, grade: false, solution: true -->
```python
def gen_code_keys(book, plain_text):
    mydict = {}
    for char in plain_text:
        if char not in mydict:
            ind =  book.find(char)
            mydict[char] = str(ind)
    return mydict
```

Output:
```
{'n': '1', 'o': '7', ' ': '4', 'i': '13', 's': '26'}
```

<!-- grade_id: cell-b929fd579a140d6c, points: 5, grade: true, solution: false -->
```python
assert gen_code_keys("Once upon a time, in a house in a land far away,", "no is no" ) == {'n': '1', 'o': '7', ' ': '4', 'i': '13', 's': '26'}
# hidden tests here
```

### Part B:  Encode a Message (5 points)
Use your `gen_code_keys(book, plain_text)` from Part A to write a function `encoder(book, plain_text)` that produces the cipher text.
Each encoded index (string) should be separated by an asterisk *.

For example: If plain_text were "no is no" and book started with "Once upon a time, in a house in a land far away,", then 
encoder(plain_text) = '1\*7\*4\*13\*26\*4\*1\*7'

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def encoder(book, plain_text):

    def gen_code_keys(book, plain_text):
        mydict = {}
        for char in plain_text:
            if char not in mydict:
                ind =  book.find(char)
                mydict[char] = str(ind)
        return mydict
    
    thedict = gen_code_keys(book, plain_text)
    res = ""

    for char in plain_text:
        res = res + '*' + thedict[char]

    return res[1:]
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 5, grade: true, solution: false -->
```python
# hidden tests here
```

### Part C:  Decode a Message (5 points)
Write a function `decoder(book, cipher_text)` that deciphers the `cipher_text` back to its original text message based on the `book`

For example: If `cipher_text were "1\*7\*4\*13\*26\*4\*1\*7" and book started with "Once upon a time, in a house in a land far away,", then 
original_text = decoder(book, cipher_text)

<!-- grade_id: cell-6a08947d8e1b5f34, points: 0, grade: false, solution: true -->
```python
def decoder(book, cipher_text):
    res = ""
    chars = cipher_text.split('*')  
    for char in chars:
        res = res + book[int(char)]  
    return res
```

<!-- grade_id: cell-90eed99978951c24, points: 5, grade: true, solution: false -->
```python
assert decoder("Once upon a time, in a house in a land far away,",  "1*7*4*13*26*4*1*7") == 'no is no'
# hidden tests here
```

### Part D:  Decode a Message (5 points)
If a character in plain_text does not appear in the book, your encoded message will decode incorrectly.
To ensure that every character in the plain text can be located, you can create a safe version of the book by appending the entire plain_text to the end of the book before encoding. This guarantees that every character in plain_text has at least one valid index in the modified "safer" book.

Write a function `gen_safe_code_keys(book, plain_text)` that generates encoding keys safely using this method.


<!-- grade_id: cell-b365b3c5ea59788f, points: 0, grade: false, solution: true -->
```python
def gen_safe_code_keys(book, plain_text):
    mydict = {}
    book = book + plain_text
    for char in plain_text:
        if char not in mydict:
            ind =  book.find(char)
            mydict[char] = str(ind)
    return mydict
```

<!-- grade_id: cell-ef211df3043fe8cc, points: 5, grade: true, solution: false -->
```python
book = "Once upon a time, in a house in a land far away,"
plain_text = "xyz"
gen_safe_code_keys(book, plain_text) == {'x': '48', 'y': '46', 'z': '50'}

# hidden tests here
```

## Problem 2 (10 pts)

Given a string s, find the first non-repeating character and return its index. If it does not exist, return -1.

**Note:**
- A character is considered non-repeating if it appears exactly once in the string.
- The comparison is case-sensitive, meaning 'A' and 'a' are treated as different characters.

<!-- grade_id: cell-6192f39ec893f265, points: 0, grade: false, solution: true -->
```python
def first_nonrepeat_char(s):
    """
    Find the first non-repeating character in a string and return its index.
    If no non-repeating character exists, return -1.

    :param s: Input string
    :return: Index of the first non-repeating character or -1
    """
    mydict = {}
    for char in s:
        mydict[char] = mydict.get(char, 0 ) + 1
    for i, char in enumerate(s):
        if mydict[char] == 1:
            return i
    return -1
```

<!-- grade_id: cell-e77cdd3ef25cfb44, points: 10, grade: true, solution: false -->
```python
# hidden tests here

```

## Problem 3 (10 pts)
Write a function that takes a string and return a dictionary. Your function will need to return the dictionary whose keys are the unique **letters** and the values are the frequency of the **letters** from the string.

**Notes:**
- The function must ignore case (treat uppercase and lowercase letters as the same).
- Non-alphabetic characters such as digits, spaces, and punctuation must not be counted.
- The keys in the dictionary should be lowercase letters only.

Example: 
Sample input string : 'W3resource!'
return: {'w': 1, 'r': 2, 'e': 2, 's': 1, 'o': 1, 'u': 1, 'c': 1}

<!-- grade_id: cell-09e5f93a7fb96627, points: 0, grade: false, solution: true -->
```python
import string
def count_letters_frequency(input_string):
    input_string = input_string.lower()
    table = '0123456789 ' + string.punctuation
    input_string = input_string.translate(str.maketrans("","",table))
    letter_count = {}
    for char in input_string:
        letter_count[char] = letter_count.get(char, 0) + 1
    return letter_count
```

<!-- grade_id: cell-f427fc8c08ea6cb1, points: 10, grade: true, solution: false -->
```python
# hidden tests here

```

## Problem 4 (15 pts)
Write a function called `coauthors_dict` that takes no parameters. It returns a dictionary with keys (author IDs) and values (lists of coauthor IDs) based on the information in the file `CA-GrQc.csv`.

<pre>
coauthors_dict() -> {13: [7596, 11196, 19170], 14: [14171], ...}
</pre>

**Note**: you need to convert the number from `str` type to `int` type, since reading from a file always produces objects of type `str`.

**Hint**: notice that the data contains errors. For example, the data shows that author 13 coauthored with himself/herself, which is meaningless. Your code should exclude oneself in one's list of coauthors.

<!-- grade_id: cell-eb9f4e347802869a, points: 0, grade: false, solution: true -->
```python
def coauthors_dict():
    coauthors = {}
    with open('CA-GrQc.csv', 'r') as file:
        counter = 0 
        for line in file:
            counter += 1
            if counter >= 5:
                authors = line.strip().split(',')
                if len(authors) == 2:
                    author1 = int(authors[0])
                    author2 = int(authors[1])
                            
                    if author1 == author2:
                        continue
                            
                    if author1 not in coauthors:
                        coauthors[author1] = []
                    if author2 not in coauthors[author1]:
                        coauthors[author1].append(author2)
                            
                    if author2 not in coauthors:
                        coauthors[author2] = []
                    if author1 not in coauthors[author2]:
                        coauthors[author2].append(author1)
    return coauthors
```

<!-- grade_id: cell-f860fb8aeb7243c0, points: 15, grade: true, solution: false -->
```python
# hidden tests here

```

## Problem 5 (10 points)
Given an array of strings strs, group the anagrams together. You can return the answer in any order.


- Example 1:

    - Input: strs = ["eat","tea","tan","ate","nat","bat"]

    - Output: [['eat', 'tea', 'ate'], ['tan', 'nat'], ['bat']] or any other correct variation of this 

    - Explanation: There is no string in strs that can be rearranged to form "bat". The strings "nat" and "tan" are anagrams as they can be rearranged to form each other. The strings "ate", "eat", and "tea" are anagrams as they can be rearranged to form each other.


- Example 2:

    - Input: strs = [""]

    - Output: [[""]]


- Example 3:

    - Input: strs = ["a"]

    - Output: [["a"]]

 

<!-- grade_id: cell-e257287828024400, points: 0, grade: false, solution: true -->
```python
def group_anagrams(strs):
    anagram_groups = {}
    
    for s in strs:
        key = ''.join(sorted(s))
        
        # Add to the appropriate group
        if key not in anagram_groups:
            anagram_groups[key] = []
        anagram_groups[key].append(s)
    
    return list(anagram_groups.values())
```

<!-- grade_id: cell-9fc16c0e3fb3c3dd, points: 10, grade: true, solution: false -->
```python
# hidden tests here

```

## Problem 6 (10 points)

In the file questions.txt, there are 9095 questions from the Stanford Question Answering Dataset, whose question words are in the list ['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old'].

Write a function named count_question_words() that reads the file and computes the frequency of each question word. The function should return a dictionary where:

    The keys are the question words.
    The values are the frequencies (type int) of those question words .

For instance, count_question_words() might return {'when': 736,
                                 'who': 984,
                                 'what': 5813,
                                 'where': 429,
                                 'how_many': 522,
                                 'how': 459,
                                 'how_much': 141,
                                 'how_old': 11}

Important: The keys in the dictionary must appear in the same order as the first occurrence of each question word in the file.

Note: For questions like “How old are you?”, treat the question word as "how_old" rather than "how".

<!-- grade_id: cell-3ec15fde7be1ee83, points: 0, grade: false, solution: true -->
```python
def count_question_words():
    #Do NOT change or delete the following 3 lines of code:
    question_words = ['what', 'who', 'when', 'where', 'how', 'how_many', 'how_much', 'how_old']
    counts = {}
    handle = open('questions.txt', 'r', encoding='UTF-8')

    first_seen_order = []
    
    for line in handle:
        words = line.lower().split()
        
        if words[0] != "how":
            qword = words[0]
        else:
            if len(words) > 1 and words[1] == "many":
                qword = "how_many"
            elif len(words) > 1 and words[1] == "much":
                qword = "how_much"
            elif len(words) > 1 and words[1] == "old":
                qword = "how_old"
            else:
                qword = "how"

        if qword in question_words:
            counts[qword] = counts.get(qword, 0) + 1
            if qword not in first_seen_order:
                first_seen_order.append(qword)

    result = {q: counts[q] for q in first_seen_order}

    return result
```

<!-- grade_id: cell-3aa499fd6c2e571a, points: 10, grade: true, solution: false -->
```python
#hidden tests here (the test cases are not shown to you)

```