Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Qinwen Zheng"
COLLABORATORS = ""
```

---

# Quiz 5
**COMPSCI 101, Duke Kunshan University**

## Problem 1

Write a function that takes a list and convert the dictionary into a list of lists and return it.

example: Original Dictionary:
{1: 'red', 2: 'green', 3: 'black', 4: 'white', 5: 'black'}

Convert the said dictionary into a list of lists:
[[1, 'red'], [2, 'green'], [3, 'black'], [4, 'white'], [5, 'black']]

<!-- grade_id: cell-f9cc0e7c8f4576a3, points: 0, grade: false, solution: true -->
```python
def dict_to_list_of_lists(input_dict):
    res = []
    for keys in input_dict.keys():
        c = []
        c.append(keys)
        c.append(input_dict[keys])
        res.append(c)
    return res       
```

Output:
```
[[1, 'red'], [2, 'green'], [3, 'black'], [4, 'white'], [5, 'black']]
```

<!-- grade_id: cell-344a8c6aab6688ff, points: 10, grade: true, solution: false -->
```python
assert dict_to_list_of_lists({1:'red',2:'green',3:'black',4:'white',5:'black'})==[[1, 'red'], [2, 'green'], [3, 'black'], [4, 'white'], [5, 'black']]
# hidden tests below
```

```python

```